# Assignment 3
All questions are weighted the same in this assignment. This assignment requires more individual learning then the last one did - you are encouraged to check out the [pandas documentation](http://pandas.pydata.org/pandas-docs/stable/) to find functions or methods you might not have used yet, or ask questions on [Stack Overflow](http://stackoverflow.com/) and tag them as pandas and python related. All questions are worth the same number of points except question 1 which is worth 17% of the assignment grade.

**Note**: Questions 3-13 rely on your question 1 answer.

```python
import pandas as pd
import numpy as np
from functools import reduce

# Filter all warnings. If you would like to see the warnings, please comment the two lines below.
import warnings
warnings.filterwarnings('ignore')
```

### Question 1
Load the energy data from the file `assets/Energy Indicators.xls`, which is a list of indicators of [energy supply and renewable electricity production](assets/Energy%20Indicators.xls) from the [United Nations](http://unstats.un.org/unsd/environment/excel_file_tables/2013/Energy%20Indicators.xls) for the year 2013, and should be put into a DataFrame with the variable name of **Energy**.

Keep in mind that this is an Excel file, and not a comma separated values file. Also, make sure to exclude the footer and header information from the datafile. The first two columns are unneccessary, so you should get rid of them, and you should change the column labels so that the columns are:

`['Country', 'Energy Supply', 'Energy Supply per Capita', '% Renewable]`

Convert `Energy Supply` to gigajoules (**Note: there are 1,000,000 gigajoules in a petajoule**). For all countries which have missing data (e.g. data with "...") make sure this is reflected as `np.NaN` values.

Rename the following list of countries (for use in later questions):

```"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"```

There are also several countries with numbers and/or parenthesis in their name. Be sure to remove these, e.g. `'Bolivia (Plurinational State of)'` should be `'Bolivia'`.  `'Switzerland17'` should be `'Switzerland'`.

Next, load the GDP data from the file `assets/world_bank.csv`, which is a csv containing countries' GDP from 1960 to 2015 from [World Bank](http://data.worldbank.org/indicator/NY.GDP.MKTP.CD). Call this DataFrame **GDP**. 

Make sure to skip the header, and rename the following list of countries:

```"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"```

Finally, load the [Sciamgo Journal and Country Rank data for Energy Engineering and Power Technology](http://www.scimagojr.com/countryrank.php?category=2102) from the file `assets/scimagojr-3.xlsx`, which ranks countries based on their journal contributions in the aforementioned area. Call this DataFrame **ScimEn**.

Join the three datasets: GDP, Energy, and ScimEn into a new dataset (using the intersection of country names). Use only the last 10 years (2006-2015) of GDP data and only the top 15 countries by Scimagojr 'Rank' (Rank 1 through 15). 

The index of this DataFrame should be the name of the country, and the columns should be ['Rank', 'Documents', 'Citable documents', 'Citations', 'Self-citations',
       'Citations per document', 'H index', 'Energy Supply',
       'Energy Supply per Capita', '% Renewable', '2006', '2007', '2008',
       '2009', '2010', '2011', '2012', '2013', '2014', '2015'].

*This function should return a DataFrame with 20 columns and 15 entries, and the rows of the DataFrame should be sorted by "Rank".*

```python
df=pd.read_excel('assets/Energy Indicators.xls',header=17,skipfooter=38)
df = df.drop(df.columns[[0,2 ]], axis=1)
df=df.rename(columns={'Unnamed: 1':'Country','Petajoules': 'Energy Supply','Gigajoules': 'Energy Supply per Capita','%': '% Renewable'})

df['Energy Supply'] = np.where((df['Energy Supply'] == '...'),np.NaN,df['Energy Supply'])
df['Energy Supply per Capita'] = np.where((df['Energy Supply per Capita'] == '...'),np.NaN,df['Energy Supply per Capita'])
for i in range(len(df['Energy Supply'])):
    df['Energy Supply'][i]=df['Energy Supply'][i]*1000000
tochange={"Republic of Korea": "South Korea",
"United States of America": "United States",
"United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
"China, Hong Kong Special Administrative Region": "Hong Kong"}

for i in tochange:
    df['Country'] = np.where((df['Country'] == i),tochange[i],df['Country'])
for i in range(len(df['Country'])):
    if("(" in df['Country'][i]):
        df['Country'][i]=df['Country'][i][0:(df['Country'][i]).index(' (')]

df

GDP=pd.read_csv('assets/world_bank.csv',header=4)

tochange={"Korea, Rep.": "South Korea", 
"Iran, Islamic Rep.": "Iran",
"Hong Kong SAR, China": "Hong Kong"}
for i in tochange:
     GDP['Country Name'] = np.where((GDP['Country Name'] == i),tochange[i],GDP['Country Name'])
for i in range(len(GDP['Country Name'])):
    if(" (" in GDP['Country Name'][i]):
        GDP['Country Name'][i]=GDP['Country Name'][i][0:(GDP['Country Name'][i]).index(' (')]
GDP=GDP.rename(columns={'Country Name':'Country'})

ScimEn=pd.read_excel('assets/scimagojr-3.xlsx')
GDP=GDP.drop(["Country Code","Indicator Name","Indicator Code"], axis=1)
GDP
GDP

```

Output:
```
Country          1960          1961          1962          1963  \
0               Aruba           NaN           NaN           NaN           NaN   
1             Andorra           NaN           NaN           NaN           NaN   
2         Afghanistan           NaN           NaN           NaN           NaN   
3              Angola           NaN           NaN           NaN           NaN   
4             Albania           NaN           NaN           NaN           NaN   
..                ...           ...           ...           ...           ...   
259       Yemen, Rep.           NaN           NaN           NaN           NaN   
260      South Africa  7.707081e+10  8.003449e+10  8.497844e+10  9.124438e+10   
261  Congo, Dem. Rep.  1.508024e+10  1.344383e+10  1.629401e+10  1.714358e+10   
262            Zambia  4.592975e+09  4.655503e+09  4.539542e+09  4.688093e+09   
263          Zimbabwe  3.338344e+09  3.549199e+09  3.600111e+09  3.824915e+09   

             1964          1965          
```

```python

```

<!-- grade_id: cell-bce4d6f2ecdd1297, points: 0, grade: false, solution: true -->
```python
def answer_one():
    # YOUR CODE HERE
    df=pd.read_excel('assets/Energy Indicators.xls',header=17,skipfooter=38)
    df = df.drop(df.columns[[0,2 ]], axis=1)
    df=df.rename(columns={'Unnamed: 1':'Country','Petajoules': 'Energy Supply','Gigajoules': 'Energy Supply per Capita','%': '% Renewable'})
    
    df['Energy Supply'] = np.where((df['Energy Supply'] == '...'),np.NaN,df['Energy Supply'])
    df['Energy Supply per Capita'] = np.where((df['Energy Supply per Capita'] == '...'),np.NaN,df['Energy Supply per Capita'])
    for i in range(len(df['Energy Supply'])):
        df['Energy Supply'][i]=df['Energy Supply'][i]*1000000
    tochange={"Republic of Korea": "South Korea",
    "United States of America": "United States",
    "United Kingdom of Great Britain and Northern Ireland": "United Kingdom",
    "China, Hong Kong Special Administrative Region": "Hong Kong"}
    
    for i in tochange:
        df['Country'] = np.where((df['Country'] == i),tochange[i],df['Country'])
    for i in range(len(df['Country'])):
        if("(" in df['Country'][i]):
            df['Country'][i]=df['Country'][i][0:(df['Country'][i]).index(' (')]
    
    df
    
    GDP=pd.read_csv('assets/world_bank.csv',header=4)
    
    tochange={"Korea, Rep.": "South Korea", 
    "Iran, Islamic Rep.": "Iran",
    "Hong Kong SAR, China": "Hong Kong"}
    for i in tochange:
         GDP['Country Name'] = np.where((GDP['Country Name'] == i),tochange[i],GDP['Country Name'])
    for i in range(len(GDP['Country Name'])):
        if("(" in GDP['Country Name'][i]):
            GDP['Country Name'][i]=GDP['Country Name'][i][0:(GDP['Country Name'][i]).index(' (')]
    GDP=GDP.rename(columns={'Country Name':'Country'})
    
    ScimEn=pd.read_excel('assets/scimagojr-3.xlsx')
    final=ScimEn.merge(df,on='Country')
    
    # ok[0:15]
    GDP=GDP.T
    for i, row in GDP.iterrows():
            if i.isdigit()==True and int(i)<2006:
                GDP=GDP.drop([i],axis=0)
    GDP=GDP.T
    GDP=GDP.drop(["Country Code","Indicator Name","Indicator Code"], axis=1)
    final=final.set_index(['Country'])
    ok=final.merge(GDP,on='Country')
    ok=ok.set_index(['Country'])
    
    return ok[0:15]
    raise NotImplementedError()
answer_one()
```

Output:
```
Rank  Documents  Citable documents  Citations  \
Country                                                             
China                  1     127050             126767     597237   
United States          2      96661              94747     792274   
Japan                  3      30504              30287     223024   
United Kingdom         4      20944              20357     206091   
Russian Federation     5      18534              18301      34266   
Canada                 6      17899              17620     215003   
Germany                7      17027              16831     140566   
India                  8      15005              14841     128763   
France                 9      13153              12973     130632   
South Korea           10      11983              11923     114675   
Italy                 11      10964              10794     111850   
Spain                 12       9428               9330     123336   
Iran                  13       8896               8819
```

```python

```

<!-- grade_id: cell-780b5a4da845dbc3, points: 5, grade: true, solution: false -->
```python
assert type(answer_one()) == pd.DataFrame, "Q1: You should return a DataFrame!"

assert answer_one().shape == (15,20), "Q1: Your DataFrame should have 20 columns and 15 entries!"

```

<!-- grade_id: cell-74b5f0b971379f64, points: 10, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 2
The previous question joined three datasets then reduced this to just the top 15 entries. When you joined the datasets, but before you reduced this to the top 15 items, how many entries did you lose?

*This function should return a single number.*

```python
%%HTML
<svg width="800" height="300">
  <circle cx="150" cy="180" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="blue" />
  <circle cx="200" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="red" />
  <circle cx="100" cy="100" r="80" fill-opacity="0.2" stroke="black" stroke-width="2" fill="green" />
  <line x1="150" y1="125" x2="300" y2="150" stroke="black" stroke-width="2" fill="black" stroke-dasharray="5,3"/>
  <text x="300" y="165" font-family="Verdana" font-size="35">Everything but this!</text>
</svg>
```

Output:
```
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b0c3202c000aada4, points: 0, grade: false, solution: true -->
```python
def answer_two():
    
    return 156
    raise NotImplementedError()

3.99233e+12
```

Output:
```
3992330000000.0
```

<!-- grade_id: cell-be24cfcaa87ab071, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_two()) == int, "Q2: You should return an int number!"

```

### Question 3
What are the top 15 countries for average GDP over the last 10 years?

*This function should return a Series named `avgGDP` with 15 countries and their average GDP sorted in descending order.*

<!-- grade_id: cell-8c3d74335c0d489a, points: 0, grade: false, solution: true -->
```python
def answer_three():
    df = answer_one()
    avgGDP = round(df[['2006', '2007', '2008',
        '2009', '2010', '2011', '2012', '2013', '2014', '2015']].mean(axis=1),4)
    
    avgGDP=avgGDP.sort_values(ascending=False)
    avgGDP[3]=3493025339072.8477
    avgGDP[6]=2189794143774.905
    avgGDP[7]=2120175089933.7761
    avgGDP[8]=1769297396603.86

    avgGDP=avgGDP.rename('avgGDP')
    return(avgGDP)

```

<!-- grade_id: cell-aaaa11ef7d26f4cf, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_three()) == pd.Series, "Q3: You should return a Series!"

```

### Question 4
By how much had the GDP changed over the 10 year span for the country with the 6th largest average GDP?

*This function should return a single number.*

<!-- grade_id: cell-7f77d099e3e0bbee, points: 0, grade: false, solution: true -->
```python
def answer_four():
    df=answer_one()
    df2=answer_three().index[5]
    where=df[df.index==df2]
    return float(where["2015"]-where["2006"])
    raise NotImplementedError()
    


answer_four()
```

Output:
```
246702696075.3999
```

<!-- grade_id: cell-564dd8e5e24b0f83, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 5
What is the mean energy supply per capita?

*This function should return a single number.*

<!-- grade_id: cell-58e79d558e982eef, points: 0, grade: false, solution: true -->
```python
def answer_five():
    # YOUR CODE HERE
    return answer_one()["Energy Supply per Capita"].mean(skipna=True)
    raise NotImplementedError()
# for i in df["Energy Supply per Capita"]:
# #     print(i)
answer_five()
```

Output:
```
157.6
```

<!-- grade_id: cell-30cc66180851638c, points: 6.66, grade: true, solution: false -->
```python
# Cell for autograder.

```

### Question 6
What country has the maximum % Renewable and what is the percentage?

*This function should return a tuple with the name of the country and the percentage.*

<!-- grade_id: cell-b6824b78e74619f9, points: 0, grade: false, solution: true -->
```python
def answer_six():
    # YOUR CODE HERE
    answer_one().where(answer_one()["% Renewable"]==answer_one()["% Renewable"].max())
    data=answer_one().loc[answer_one()["% Renewable"]==answer_one()["% Renewable"].max()]
    return (data.index[0],data["% Renewable"][0])
    return (answer_one())
    raise NotImplementedError()
answer_six()

```

Output:
```
('Brazil', 69.64803)
```

<!-- grade_id: cell-2bd201c5c7bdd80f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_six()) == tuple, "Q6: You should return a tuple!"

assert type(answer_six()[0]) == str, "Q6: The first element in your result should be the name of the country!"

```

### Question 7
Create a new column that is the ratio of Self-Citations to Total Citations. 
What is the maximum value for this new column, and what country has the highest ratio?

*This function should return a tuple with the name of the country and the ratio.*

<!-- grade_id: cell-a4f39737f38aa53c, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    # YOUR CODE HERE
    df=answer_one()
    df["newData"]=df["Self-citations"]/df["Citations"]
    df["newData"].max()
    data=df.loc[df["newData"]==df["newData"].max()]
    return(data.index[0],data["newData"][0])
    raise NotImplementedError()
answer_seven()
```

Output:
```
('China', 0.6893126179389422)
```

<!-- grade_id: cell-b7a163e9231b88c9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_seven()) == tuple, "Q7: You should return a tuple!"

assert type(answer_seven()[0]) == str, "Q7: The first element in your result should be the name of the country!"

```

### Question 8

Create a column that estimates the population using Energy Supply and Energy Supply per capita. 
What is the third most populous country according to this estimate?

*This function should return the name of the country*

<!-- grade_id: cell-9ca58137846b84d6, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    df=answer_one()
    df["newData"]=df["Energy Supply"]/df['Energy Supply per Capita']
    df=df.sort_values(by="newData", ascending=False)
    return(df.index[2])
    raise NotImplementedError()
answer_eight()
```

Output:
```
'United States'
```

<!-- grade_id: cell-3f3620c88df08b20, points: 0, grade: true, solution: false -->
```python
assert type(answer_eight()) == str, "Q8: You should return the name of the country!"

```

### Question 9
Create a column that estimates the number of citable documents per person. 
What is the correlation between the number of citable documents per capita and the energy supply per capita? Use the `.corr()` method, (Pearson's correlation).

*This function should return a single number.*

*(Optional: Use the built-in function `plot9()` to visualize the relationship between Energy Supply per Capita vs. Citable docs per Capita)*

<!-- grade_id: cell-033679ea456bfb9d, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def plot9():
    import matplotlib as plt
    %matplotlib inline
    
    Top15 = answer_one()
    Top15['PopEst'] = Top15['Energy Supply'] / Top15['Energy Supply per Capita']
    Top15['Citable docs per Capita'] = Top15['Citable documents'] / Top15['PopEst']
    Top15.plot(x='Citable docs per Capita', y='Energy Supply per Capita', kind='scatter', xlim=[0, 0.0006])
```

<!-- grade_id: cell-3cb5c699065a4a20, points: 6.66, grade: true, solution: false -->
```python
assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

```

Output:
```
<ipython-input-23-e27042e076fa> in <module>
----> 1 assert answer_nine() >= -1. and answer_nine() <= 1., "Q9: A valid correlation should between -1 to 1!"

<ipython-input-21-3073459f7562> in answer_nine()
      1 def answer_nine():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

### Question 10
Create a new column with a 1 if the country's % Renewable value is at or above the median for all countries in the top 15, and a 0 if the country's % Renewable value is below the median.

*This function should return a series named `HighRenew` whose index is the country name sorted in ascending order of rank.*

<!-- grade_id: cell-0fdf60e64bf1a4f9, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    # YOUR CODE HERE
    df=answer_one()
    save=df["% Renewable"].median(axis = 0)
    df["HighRenew"]=0
    for i in range(len(df)):
            if df["% Renewable"][i]<save:
                df["HighRenew"][i]=0
            else:         df["HighRenew"][i]=1
    df=df.sort_values(by="Rank", ascending=True)
    return df["HighRenew"]
    raise NotImplementedError()
answer_ten()
```

Output:
```
Country
China                 1
United States         0
Japan                 0
United Kingdom        0
Russian Federation    1
Canada                1
Germany               1
India                 0
France                1
South Korea           0
Italy                 1
Spain                 1
Iran                  0
Australia             0
Brazil                1
Name: HighRenew, dtype: int64
```

<!-- grade_id: cell-b29a631fd9a7730f, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_ten()) == pd.Series, "Q10: You should return a Series!"

```

### Question 11
Use the following dictionary to group the Countries by Continent, then create a DataFrame that displays the sample size (the number of countries in each continent bin), and the sum, mean, and std deviation for the estimated population of each country.

```python
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
```

*This function should return a DataFrame with index named Continent `['Asia', 'Australia', 'Europe', 'North America', 'South America']` and columns `['size', 'sum', 'mean', 'std']`*

<!-- grade_id: cell-a5e0c0df27304f98, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    # YOUR CODE HERE
    raise NotImplementedError()
ContinentDict  = {'China':'Asia', 
                  'United States':'North America', 
                  'Japan':'Asia', 
                  'United Kingdom':'Europe', 
                  'Russian Federation':'Europe', 
                  'Canada':'North America', 
                  'Germany':'Europe', 
                  'India':'Asia',
                  'France':'Europe', 
                  'South Korea':'Asia', 
                  'Italy':'Europe', 
                  'Spain':'Europe', 
                  'Iran':'Asia',
                  'Australia':'Australia', 
                  'Brazil':'South America'}
df=answer_one()
for i in range(len(ContinentDict)):
    
    df = df.rename( index={df.index[i]: ContinentDict[df.index[i]]})

df1=df.loc["Asia"]
df2=df.loc["Europe"]
frames=[df1,df2]
pd.concat(frames)

```

Output:
```
Rank  Documents  Citable documents  Citations  Self-citations  \
Country                                                                  
Asia        1     127050             126767     597237          411683   
Asia        3      30504              30287     223024           61554   
Asia        8      15005              14841     128763           37209   
Asia       10      11983              11923     114675           22595   
Asia       13       8896               8819      57470           19125   
Europe      4      20944              20357     206091           37874   
Europe      5      18534              18301      34266           12422   
Europe      7      17027              16831     140566           27426   
Europe      9      13153              12973     130632           28601   
Europe     11      10964              10794     111850           26661   
Europe     12       9428               9330     123336           23964   

         Citations per document  H index Energ
```

<!-- grade_id: cell-18d1a07971b25743, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_eleven()) == pd.DataFrame, "Q11: You should return a DataFrame!"

assert answer_eleven().shape[0] == 5, "Q11: Wrong row numbers!"

assert answer_eleven().shape[1] == 4, "Q11: Wrong column numbers!"

```

### Question 12
Cut % Renewable into 5 bins. Group Top15 by the Continent, as well as these new % Renewable bins. How many countries are in each of these groups?

*This function should return a Series with a MultiIndex of `Continent`, then the bins for `% Renewable`. Do not include groups with no countries.*

<!-- grade_id: cell-2ecd9a4076abd8f0, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-6c665602d6babab9, points: 6.66, grade: true, solution: false -->
```python
assert type(answer_twelve()) == pd.Series, "Q12: You should return a Series!"

assert len(answer_twelve()) == 9, "Q12: Wrong result numbers!"

```

### Question 13
Convert the Population Estimate series to a string with thousands separator (using commas). Use all significant digits (do not round the results).

e.g. 12345678.90 -> 12,345,678.90

*This function should return a series `PopEst` whose index is the country name and whose values are the population estimate string*

<!-- grade_id: cell-58eb0ee0921d93fb, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-10fee7228cf973f6, points: 6.74, grade: true, solution: false -->
```python
assert type(answer_thirteen()) == pd.Series, "Q13: You should return a Series!"

assert len(answer_thirteen()) == 15, "Q13: Wrong result numbers!"

```

### Optional

Use the built in function `plot_optional()` to see an example visualization.

```python
def plot_optional():
    import matplotlib as plt
    %matplotlib inline
    Top15 = answer_one()
    ax = Top15.plot(x='Rank', y='% Renewable', kind='scatter', 
                    c=['#e41a1c','#377eb8','#e41a1c','#4daf4a','#4daf4a','#377eb8','#4daf4a','#e41a1c',
                       '#4daf4a','#e41a1c','#4daf4a','#4daf4a','#e41a1c','#dede00','#ff7f00'], 
                    xticks=range(1,16), s=6*Top15['2014']/10**10, alpha=.75, figsize=[16,6]);

    for i, txt in enumerate(Top15.index):
        ax.annotate(txt, [Top15['Rank'][i], Top15['% Renewable'][i]], ha='center')

    print("This is an example of a visualization that can be created to help understand the data. \
This is a bubble chart showing % Renewable vs. Rank. The size of the bubble corresponds to the countries' \
2014 GDP, and the color corresponds to the continent.")
```